```python
# Initialize Otter
import otter
grader = otter.Notebook("lab06.ipynb")
```

# Lab 06: Writing and Cracking the LFSR with a Plaintext Attack 
In this week’s lab, you’ll learn how to crack the LFSR cipher using a plaintext attack.

Let’s get started!

# Part I: Dictionaries

In Python, we can use a data structure called a dictionary to convert one string into another. Run the cell below to load a dictionary that converts upper case Roman characters into 5-bit binary strings. This may look familiar from the story of the Princess and the Rats.

You may recall that there was a problem with this code. A 5-bit binary code requires a 32 character alphabet. We need an additional 6 characters. In this case, we will use: a comma, a period, a question mark, an exclamation point, an @ symbol (because it resembles a rat viewed from above), and finally a hashtag (because it looks like a cage for keeping rats).

```python
binary = {"A":"00000","B":"00001", "C":"00010", "D":"00011", "E":"00100", "F":"00101", "G":"00110", "H":"00111", "I":"01000", "J":"01001", "K":"01010", "L":"01011", "M":"01100", "N":"01101", "O":"01110","P":"01111", "Q":"10000", "R":"10001", "S":"10010", "T":"10011", "U":"10100", "V":"10101", "W":"10110", "X":"10111", "Y":"11000","Z":"11001", ",":"11010",".":"11011","?":"11100","!":"11101","@":"11110","#":"11111"}

```

To use this dictionary, all we need to do is use the name of the dictionary and square brackets around the string we would like to "translate".  

```python
print(binary["A"])
```

You might think that there is some clever and efficient way to "reverse" this dictionary so that you don't have to enter in a long stream of string equivalencies. If so, you are correct!

All we need to do is define a new dictionary, this time called "roman_char" that will be the reverse of our dictionary, binary, above.

```python
roman_char = dict([[c,b] for b,c in binary.items()])

#this command will print the roman character that we have equated to 30 

print(roman_char["11110"])
```

Use the dictionary to write a function that will translate a short message such as "Flee Rat" into binary. But first, run the cell below.

```python
def text_clean( text, LETTERS = 'ABCDEFGHIJKLMNOPQRSTUVWXYZ,.?!@#'):
    """
    Arguments:
        text (str): a piece of text for cleaning
        LETTERS (str, optional): defines the alphabet of allowable characters
    Returns:
        (str): text with only the characters also found in LETTERS
               lower-case letters in text will be made upper-case  
    """
    cleaned_text = '' 
    
    for character in text: 
        if character.upper() in LETTERS:
            cleaned_text += character.upper()
    
    return cleaned_text

```

```python
def binary_convert(text, LETTERS = "ABCDEFGHIJKLMNOPQRSTUVWXYZ,.?!@#"):
    message = ...
    ...
```

```python
grader.check("q1_1")
```

# Part II: Recovering our plaintext message

Now that we can put a message into binary, we can also put a binary message back into Roman characters. But this time, there is one little wrinkle. We will have to break our binary message into short strings of 5 1s and 0s before feeding it into the dictionary.

In the cell below, write a function to break a binary message into chunks of 5 and translate each one.

If you do this correctly, you will be able to read this message, the opening lines of one of the oldest national anthems in the world, recently updated: "001100111000011100100000010101001000011010001001000000010011000100011100000100 0101011001001001001110101001000101010010000110100110110100101101110011010011001 0110100010101001000111010100100010110101110000010101100100010100100001101001101 1010"


```python
def roman_convert(text, LETTERS = "ABCDEFGHIJKLMNOPQRSTUVWXYZ,.?!@#"):
    message = ...
    output = ...
    ...
        p = ...
        ...
    output = ...
    ...

```

```python
grader.check("q2_1")
```

<!-- BEGIN QUESTION -->

Type your translation of the binary stream above into the cell below as a string named "solution".

```python

```

<!-- END QUESTION -->

# Part III: XORing two binary streams together

Unfortunately, simply writing a message in binary characters is not a very good way to send a secret message. First of all, the resulting message is very long. Secondly, it is not hard to guess that A = 00000 or that B = 00001, etc. We could scramble the binary equivalencies, but then we would have something that was no different from a Caesar cipher with all it's attendant vulnerabilities. So, to get around this problem, we can XOR our binary stream with some other binary stream to encipher it. Let's use the key 01001.

Try running the cell below to see how this works.

```python
a = 10100
b = 10010

print(a^b)
```

Immediately, you see that the result should have been 00110, but the leading zeros were cut off. Even worse, remove the hashtags to see the error message that appears when you run the cell below.

```python
#a = "10100"
#b = "10010"

#print(a^b)
```

Finally, sometimes you can get <i>very</i> strange results. Try running the cell below. 

```python
a = 11010
b = 10011

print (a^b)
```

Instead, try this: the two binary streams start off as strings. Then their integer forms are XOR'd together. The "comma 2" indicates that the strings are binary representations. However, notice that the leading 0 zeros are still missing. In your function, you will have to make certain that the length of the resulting string is the same as the lengths of the string that you are enciphering.

You will also have to repeat the keystring until it is as long as the message.

```python
a="11011111101100110110011001011101000"
b="11001011101100111000011100001100001"
y=int(a,2) ^ int(b,2)
print('{0:b}'.format(y))
```

```python
def xor_function(messagestring, keystring):
    ...
                          
```

```python
grader.check("q4_1")
```

# Part IV: Putting it All Together and Encrypting Our First Message.

Having a short and simple key is useful because it would be easy to securely share that key with your allies. Let's use the key "01000" to encrypt the message "Lancashire". If it works correctly, you will get <b>DIFKI,PAZM</b>.

Write a function that converts the message into binary, XOR's it with the key, and then puts it into Roman characters again for ease of transmission. Since the XOR function can encipher as well as decipher, you can also make the function work backwards. Remember our convention that ciphertext is in upper case and plaintext is in lower case. Make sure that you repeat the key enough times that it "covers" the message.

If you do this correctly, then you will be able to decipher and read this message: ",HJOJ#YH@NP.?LCS!YWBH,UBSEMXB#YH@NEE!COWBH,UMSKY??OPOOWBH,U". It is the opening lines of one of the oldest national anthems in the world, recently updated. 


```python
def xor_encrypt(message, key, encrypt = True, LETTERS = "ABCDEFGHIJKLMNOPQRSTUVWXYZ,.?!@#"):
    ...

```

```python
grader.check("q4_2")
```

Interestingly, a message encrypted this way with a short key will resemble our caesar cipher quite closely and have some of the same vulnerabilities. Use the function below to experiment with several short messages and 5-bit keys. In particular, try a 5-bit key and then use the Caesar cipher key that corresponds to that same number. For example, try "01100" and 12.

```python
def caesar(message, key, LETTERS = "ABCDEFGHIJKLMNOPQRSTUVWXYZ,.?!@#"):
    message = text_clean(message)
    output = ""

    for char in message:
      index = LETTERS.find(char)
      newindex = (index + key) % len(LETTERS)
      output = output + (LETTERS[newindex])
    return(output)
```

What do you notice? Write a few sentences in the cell below.

<!-- BEGIN QUESTION -->



```python

```

<!-- END QUESTION -->

# Part VI: Creating an LFSR keystream

The LFSR function lets us have a short and simple key and yet, generate a keystream that is much more complex than that short key would initally allow.

In this example, we will start with a 5-bit seed. The "rule" we use will be that the left most bit in the table will be the XOR sum of the first, second, and fourth bits when the seed is read from right to left. As usual, the keystream will be the outputs of the right most column. However, an LFSR can have many more digit in its seed and can have any rule.

Using a seed of 11001 (as a string), the first 10 digits of the keystream should be 1001101110. Create a function that takes in the 5-bit seed and outputs a keystream of any integer length.

```python
def LFSR_gen(seed, length):
    ...
```

```python
grader.check("q6")
```

# Putting it All Together: The Reboot

Now we can use our new and improved pseudorandom keystream to encipher our message. Because the XOR function is exactly the same deciphering and enciphering, we can write a function with encryption set to the default value of true that will encrypt if not otherwise specified. Then we can create a function that will take in a message and a 5-bit seed and return plaintext to us if ciphertext is provided, and also do the reverse. As you write your function, recall our convention of putting plaintext in lower case and ciphertext in upper case.

If you succeed, "Whenever You Need Somebody", the title of a popular pop song, should become FJG@KXX#,!,PXKBBAOXPMQW. Also, if you input FJG@KXX#,!,PXKBBAOXPMQW, then your output should be wheneveryouneedsomebody. Notice that we lost the capitals in the title and also the spaces, but the message is easy to understand regardless. 


```python
def LFSR_encipher(message, seed, encipher = True, LETTERS = "ABCDEFGHIJKLMNOPQRSTUVWXYZ,.?!@#"):
    ...

```

```python
grader.check("q7")
```

<!-- BEGIN QUESTION -->

# Your Agents Have Intercepted a Secret Message From A Powerful Global Organization

This is the message:
    
@KXX#E!DPTIKGK,!,W?DGGKTVAP@OJX!,!,B!YP@KXX#E!DPT#W@OT!,PQOPQNGBKTAWMHDGGKTVAP@OOTEGLAWR#,@KXX#E!DPT?CLIM!NDLKPX.GCIM@DCAKJYOJ.KC@NFH#RLAW

Your agents also believe that the message was enciphered with the same LFSR keystream our program can generate and a seed of "11001". What a lucky coincidence! Decipher the message. What horrible thing have your enemies done to you? Answer in the cell below.

```python

```

<!-- END QUESTION -->

Submitting your work 
You’re done with this Lab! All assignments in the course will be distributed as notebooks like this one, and you will submit your work by doing the following:

Save your notebook
Restart the kernel and run up to this cell.
Run all the tests by running the cell containing grader.check_all(). Make sure they pass the way you expect them to.
Run the cell below with the code grader.export(...).
Download the file named labXX.zip, found in the explorer pane on the left side of the screen.
Upload labXX-<date-time stamp>.zip to the corresponding lab assignment on Canvas.

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit.

When done exporting, find the `.zip` file in the left side of the screen in the file browser, right-click, and select **Download**. You'll submit this `.zip` file for the assignment Gradescope for grading.

```python
grader.export(pdf=False, force_save=True)
```

Tests for q1_1 (points: None)
```
>>> binary_convert("b") != None
True
```
failure_message: Check that you are returning the ciphertext, not printing it
```
>>> binary_convert("binary")
'000010100001101000001000111000'
```


Tests for q2_1 (points: None)
```
>>> roman_convert("10000") != None
True
```
failure_message: Check that you are returning the ciphertext, not printing it
```
>>> roman_convert("10001")
'r'
```
failure_message: Check that you are returning the ciphertext, not printing it
```
>>> roman_convert("11111")
'#'
```
failure_message: Check that you are using our expanded alphabet correctly and that you are returning plaintext in lower case.


Tests for q4_1 (points: None)
```
>>> xor_function("1000101001010100101001", "10010") != None
'0001111011110001100011'
```
failure_message: Check that you are returning and not printing
```
>>> xor_function("1000101001010100101001", "10010")
'0001111011110001100011'
```
failure_message: Check that you are inputting two strings


Tests for q4_2 (points: None)
```
>>> xor_encrypt("1000101001010100101001", "10010") != None
```
failure_message: Check that you are returning and not printing
```
>>> xor_encrypt("helloworld", "10010")
'VWZZ?E?DZR'
```
failure_message: Check that you are xor'ing with a keystream that is long enough


Tests for q6 (points: None)
```
>>> LFSR_gen("11100", 10) != None
True
```
failure_message: Check that you are returning and not printing
```
>>> LFSR_gen("11100", 15)
'001111010110010'
```
failure_message: Temporarily print each row to make sure it is doing what you think it should be doing.


Tests for q7 (points: None)
```
>>> LFSR_encipher("Hello World","00011") != None
True
```
failure_message: Check that are returning and not printing.
```
>>> LFSR_encipher("HelloWorld","00011")
'#OGTE.W.G.'
```
failure_message: Check that are feeding two strings into your function.
```
>>> LFSR_encipher("Hell0W0rld","00011")
'#OGT??TJ'
```
failure_message: Check that you are using text_clean to clean your function.
```
>>> LFSR_encipher("#OGTE.W.G.","00011", False)
'helloworld'
```
failure_message: Check that your function can decipher as well as encipher and returns plaintext in lower case.
