```python
# Initialize Otter
import otter
grader = otter.Notebook("lab02.ipynb")
```

# Lab 02: Caesar Cipher

In today's lab, you'll learn how to:

1. implement the Caesar cipher
2. implement the Caesar cipher with other alphabets
3. read text from a file
4. clean and format text

This lab covers parts of [Chapter 4](https://macs4200.org/chapters/04/2/caesar-cipher.html) and [Chapter 5](https://macs4200.org/chapters/05/3/file-io.html) of the online textbook. Let's get started!

## Part 1: Caesar Cipher

So far we've seen the Caesar cipher described mathematically as:

$$ C \equiv P + k \pmod m$$

where:
* $C$ is a ciphertext character represented as an integer
* $P$ is a plaintext character represented as an integer
* $k$ is an integer key
* $m$ is the number of characters in the used alphabet

### Question 1.1

In the cell below, you'll find the set up for completing a Caesar cipher using the standard 26 letter alphabet. To start:

* Choose a plaintext that contains a single character
* choose a valid key for the Caesar cipher

Then, complete the code so the plaintext is converted to ciphertext using the Caesar cipher algorithm. You should convert the plaintext character to an integer, apply the key to obtain an integer for the ciphertext, and then convert the ciphertext integer to a character.

**Hint**: Using the `.find` method and string indexing will be helpful for a few of these steps!

```python
LETTERS = 'ABCDEFGHIJKLMNOPQRSTUVWXYZ'
plaintext = ...
key = ...

plaintext_numerical = ...
ciphertext_numerical = ...
ciphertext = ...
print(ciphertext)
```

```python
grader.check("q1_1")
```

### `.upper` and `.replace`

In the next few questions we'll be updating the code so it will work on string longer than 1 character. However, as soon as we start inputting text as strings it could easily open the door to unpredictability. Will the plaintext be all uppercase? All lowercase? Will it contain spaces, numbers, or other characters? These things make a difference because if a computer is expecting to use an uppercase letter but it gets a space or lowercase letter, your code will likely not work. We will need to take care to "sanitize" our plaintext to ensure the string we're using is predictably formatted.


One quick way to clean your string is to make all characters uppercase using the `.upper` method. Another way to is look for common characters someone may include, but are not part of our 26 character alphabet, and replace them. There is a string method named `.replace` that can do this for you.

Run the cell below to see `.replace` in action.

```python
message = 'Hello, my friend! How are you? I am well'
partially_cleaned = message.replace(' ', '')
print( partially_cleaned )
```

You can see that `.replace` works by attaching it to the end of a string (or a variable that contains a string) and providing two pieces of information: the character you want to replace (in this case a space, `' '`) and what you want to replace it with (in this case an empty string `''`).

You can also chain together two `.replace` calls if you need to replace more than one character from a message.

```python
more_cleaned = message.replace(' ', '').replace('!', '')
print( more_cleaned )
```

You can also combine `.replace` with other string methods, like `.upper`

```python
print( message.replace(' ', '').replace('!', '').upper() )
```

This is a simple way to take a plaintext message and prepare it for use in a cipher, since our alphabet is currently defined to be the characters A-Z (all uppercase). You'll notice the code above would *not* have removed any other punctuation besides the exclamation point, so it's not perfect. It's actually **really** difficult to explain all the different characters you don't want to work with. Later on in the course we'll cover a better way to prepare text that allows us to explain what characters we **do** want to work with, and any character not on the list won't be kept.

For now, assume that all messages will only contain uppercase letters, lowercase letters, spaces, commas, question marks, and periods. Before being passed onto to be encrypted, all characters should be made uppercase, and other characters removed. Don't worry about numbers or other punctuation right now.

### Question 1.2

In the cell below, use `.replace` and `.upper` to fully prepare the string stored to `message` to be encrypted. Save the cleaned string to `fully_cleaned`

```python
fully_cleaned = ...
print(fully_cleaned)
```

```python
grader.check("q1_2")
```

### Question 1.3

Now you're ready to start enciphering messages that are longer than 1 character.

Below you'll find the plaintext message, "What is the answer, you ask? I will tell you!" stored to the string `plaintext`. You should first "clean" this message as discussed in the previous question. A key has already been chosen for you.

A `for` loop has already been started that will iterate over the string `clean_plaintext` one character at a time. You should complete the lines of code inside the loop so that when the loop finishes, all the characters of the ciphertext are stored in the string `ciphertext`. The last line of code will print it to the screen to help you verify it ran correct.

```python
LETTERS = 'ABCDEFGHIJKLMNOPQRSTUVWXYZ'
plaintext = 'What is the answer, you ask? I will tell you!'
clean_plaintext = ...
key = 15

# Creating an empty string that will eventually store your ciphertext
ciphertext = ''

for char in clean_plaintext:
    plaintext_numerical = ...
    ciphertext_numerical = ...
    ciphertext = ...
    
print(ciphertext)
```

```python
grader.check("q1_3")
```

At this point you've got some great code that is enciphering messages like a champion! Let's see how it does deciphering a message. But let's make it interesting! Instead of a short message, let's see how you can read in a long message from a file.

### File Input / Output

In Python, you can read in text from a simple text file that's generated by word processors. These types of files are not fancy, you can not have different fonts, font sizes, images, or the like. They just contain the characters. Most word processors like Word and Google Docs have an option to save your documents as `.txt` files, but you can also use Notepad (on Windows) or Text Editor (Mac) to easily make these files.

Assuming you have such a file in the same folder as your notebook, you can read the text it contains into your notebook as a string and assign it to a variable. Alongside this lab notebook you'll find a file named `caesar-ciphertext.txt`. Here's the command to read it in as a string named `caesar_ciphertext`.

```python
with open('caesar-ciphertext.txt') as f:
    caesar_ciphertext = f.read()
    
print(caesar_ciphertext)
```

You can see that this is an encrypted message that has only capital letters and spaces to group the characters into blocks of length 5. We'll attempt to decipher this message.

### Question 1.4

In the cell below you'll find similar starter code as the previous problem (the words plaintext and ciphertext have been swapped to indicate the change in our goal). The ciphertext from the file has been cleaned by removing the spaces.

You should make an initial guess for the key, then complete the loop below to decipher the message. You should be able to reuse your code from an earlier problem, and just make a few minor tweaks.

It's unlikely that you'll guess the key on the first attempt, so keep trying keys until you have successfully deciphered the message.

**Note:** Don't forget, plaintext letters should be lowercase!

_Type your answer here, replacing this text._

```python
LETTERS = 'ABCDEFGHIJKLMNOPQRSTUVWXYZ'

# Remove all the spaces
clean_ciphertext = caesar_ciphertext.replace(' ', '')

key = ...

# Creating an empty string that will eventually store your ciphertext
plaintext = ''

for char in clean_ciphertext:
    ciphertext_numerical = ...
    plaintext_numerical = ...
    plaintext = ...
    
print(plaintext)
```

```python
grader.check("q1_4")
```

You can also save a string to a text file, allowing you to send it by email, upload the file to Google Drive, etc. The code to save a string stored in `plaintext` to a file named `caesar-plaintext.txt` is shown below.

```python
with open('caesar-plaintext.txt', 'w') as f:
    print(plaintext, file=f, end='')
```

You should see `caesar-plaintext.txt` appear in your folder after a few moments. For more details on all the configuration options for writing to a file you can read [Chapter 5, Section 3](https://macs4200.org/chapters/05/3/file-io.html).

# Submitting your work
You're done with this Lab! All assignments in the course will be distributed as notebooks like this one, and you will submit your work by doing the following:
* Save your notebook
* Restart the kernel and run up to this cell.
* Run all the tests by running the cell containing `grader.check_all()`. Make sure they pass the way you expect them to.
* Run the cell below with the code `grader.export(...)`.
* Download the file named `labXX.zip`, found in the explorer pane on the left side of the screen.
* Upload `labXX-<date-time stamp>.zip` to the corresponding lab assignment on Canvas.

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit.

When done exporting, find the `.zip` file in the left side of the screen in the file browser, right-click, and select **Download**. You'll submit this `.zip` file for the assignment Gradescope for grading.

```python
grader.export(pdf=False, force_save=True)
```

Tests for q1_1 (points: 1)
```
>>> len(plaintext) == 1 
True
```
failure_message: Your string is too long
```
>>> type(key) == int
True
```
failure_message: Your key is not an integer
```
>>> plaintext_numerical != -1
True
```
failure_message: It looks like you didn't convert your plaintext to an upper case letter before using .find. Try using .upper() to ensure it can be found in LETTERS.
```
>>> LETTERS[(LETTERS.find(plaintext.upper())+key)%len(LETTERS)] == ciphertext
True
```


Tests for q1_2 (points: 1)
```
>>> '!' not in fully_cleaned
True
```
```
>>> ' ' not in fully_cleaned
True
```
```
>>> '?' not in fully_cleaned
True
```
```
>>> ',' not in fully_cleaned
True
```
```
>>> '.' not in fully_cleaned
True
```
```
>>> fully_cleaned
'HELLOMYFRIENDHOWAREYOUIAMWELL'
```


Tests for q1_3 (points: 1)
```
>>> any( bad_chars in clean_plaintext for bad_chars in ['?', '!', ' ', '.', ','] )
False
```
failure_message: You still have some punctuation in your message after cleaning
```
>>> clean_plaintext.upper() == clean_plaintext
True
```
failure_message: You didn't make the characters uppercase
```
>>> ciphertext
'LWPIXHIWTPCHLTGNDJPHZXLXAAITAANDJ'
```


Tests for q1_4 (points: 1)
```
>>> plaintext == plaintext.lower()
True
```
failure_message: You didn't make the plaintext lowercase
```
>>> len(plaintext)
1154
```
failure_message: Your message isn't the correct length
```
>>> (key * 123) % 472
203
```
failure_message: If this test fails, but you obtained readable plaintext check to make sure you are subtracting the key value, not adding the inverse key.
```
>>> plaintext
'weleftinprettygoodtimeandcameafternightfalltoklausenburghhereistoppedforthenightatthehotelroyaleihadfordinnerorrathersupperachickendoneupsomewaywithredpepperwhichwasverygoodbutthirstymemgetrecipeforminaiaskedthewaiterandhesaiditwascalledpaprikahendlandthatasitwasanationaldishishouldbeabletogetitanywherealongthecarpathiansifoundmysmatteringofgermanveryusefulhereindeedidontknowhowishouldbeabletogetonwithoutithavinghadsometimeatmydisposalwheninlondonihadvisitedthebritishmuseumandmadesearchamongthebooksandmapsinthelibraryregardingtransylvaniaithadstruckmethatsomeforeknowledgeofthecountrycouldhardlyfailtohavesomeimportanceindealingwithanoblemanofthatcountryifindthatthedistricthenamedisintheextremeeastofthecountryjustonthebordersofthreestatestransylvaniamoldaviaandbukovinainthemidstofthecarpathianmountainsoneofthewildestandleastknownportionsofeuropeiwasnotabletolightonanymaporworkgivingtheexactlocalityofthecastledraculaastherearenomapsofthiscountryasyettocomparewithourownordnancesurveymapsbutifoundthatbistritztheposttownnamedbycountdraculaisafairlywellknownplaceishallenterheresomeofmynotesastheymayrefreshmymemorywhenitalkovermytravelswithmina'
```
