```python
# Initialize Otter
import otter
grader = otter.Notebook("lab04A.ipynb")
```

# Lab 04
## The Affine Cipher

### Question 1
Complete the Affine Cipher key that we began in Webinar on August 31st

```python
def affine_encipher(m,a,message, LETTERS = "ABCDEFGHIJKLMNOPQRSTUVWXYZ"):
    message_cleaned = ...
    ciphertext = ""
    for char in message_cleaned:
        num = LETTERS.find(char)
        ...
        
    return ciphertext
```

```python
grader.check("q1_1")
```

## Question 2 ##
Modify your code sot that it uses an alphabet with both capital and lowercase letters. Use a string called BIG_LETTERS = "ABCDEFGHIJKLMNOPQRSTUVWXYZabcdefghijklmnopqrstuvwxyz"

```python
def affine_encipher_big(m, a, message, BIG_LETTERS = "ABCDEFGHIJKLMNOPQRSTUVWXYZabcdefghijklmnopqrstuvwxyz"):
    ciphertext = ""
...
```

```python
grader.check("q2_1")
```

# Question 3
Use your code to encrypt the message "A secret's worth depends on the people from whom it must be kept." Use BIG_LETTERS and a key of 11,7.

(This is from a book called <i>The Shadow of the Wind</i> by Carlos Ruiz Zafón)

Once you have finished with the encryption, save your message to a string called: zafon

```python
zafon = ...
```

```python
grader.check("q3_1")
```

# Question 4: Part I
Write a function called decipher_affine that takes in the inputs m, a, and an enciphered message string and deciphers the message. The values for m and a should be the values that were originally used to encipher the message. Keep in mind that your alphabet now has <b>52 letters</b>

How many valid values for m will it have? (Include m = 1, which makes Affine identical to Caesar, and the "stupid key": m = 1, a = 0.) Type your answer into the cell below: 

```python
number_of_valid_values = ...
```

```python
grader.check("q4_1")
```

In Activity 10, you will write or have written some "Helper Functions". One finds the greatest common divisor of two numbers and one uses that one to find the multiplicative inverse of a number. In the cells below,copy those functions. You will need to find the multiplicative inverse of the key you and your ally have agreed upon in order to decipher their message to you.

```python
def egcd(a, b):
    ...
```

```python
def mult_inverse(n, m = 26):
...
```

```python
grader.check("q4_2")
```

# Question 4: Part II
Using the functions you wrote above, you now have all that you need to easily decipher a message from an ally. Suppose you now <i>recieve</i> the message "HXZDMZiXPfMiGOZqZUOXfUiGZqZfqyZkMfJPGfJRiJtXisZnZqi" from your friend. Write a function that takes in the message and the key (which you know is [11, 7] by prior arrangement with your ally) and returns the plaintext, in other words, it can decipher this message.

```python
def big(message, m, a, BIG_LETTERS = "ABCDEFGHIJKLMNOPQRSTUVWXYZabcdefghijklmnopqrstuvwxyz"):
    plaintext = ""

```

```python
grader.check("q4_3")
```

# Submitting your work
You're done with this Lab! All assignments in the course will be distributed as notebooks like this one, and you will submit your work by doing the following:
* Save your notebook
* Restart the kernel and run up to this cell.
* Run all the tests by running the cell containing `grader.check_all()`. Make sure they pass the way you expect them to.
* Run the cell below with the code `grader.export(...)`.
* Download the file named `labXX.zip`, found in the explorer pane on the left side of the screen.
* Upload `labXX-<date-time stamp>.zip` to the corresponding lab assignment on Canvas.

```python

```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit.

When done exporting, find the `.zip` file in the left side of the screen in the file browser, right-click, and select **Download**. You'll submit this `.zip` file for the assignment Gradescope for grading.

```python
grader.export(pdf=False, force_save=True)
```

Tests for q1_1 (points: None)
```
>>> affine_encipher(3,1,"b") != None
True
```
failure_message: Check that you are returning the ciphertext, not printing it
```
>>> affine_encipher(3,1,"b")
'E'
```
failure_message: Check that your cleaned plaintext is uppercase before encrypting
```
>>> affine_encipher(3,1,"b!., ")
'E'
```
failure_message: Check that you cleaning your message free of periods, exclamation points, commas, and spaces.
```
>>> affine_encipher(19,5,"Z")
'M'
```
failure_message: Check that you're modding by the length of the alphabet (26)


Tests for q2_1 (points: None)
```
>>> affine_encipher_big(3,1,"b") != None
True
```
failure_message: Check that you are returning the ciphertext, not printing it
```
>>> affine_encipher_big(3,11,"Z")
'i'
```
failure_message: Check that you moding by the length of the larger alphabet
```
>>> affine_encipher_big(3,11,".' F")
'a'
```
failure_message: Check that you are "cleaning" the message of apostrophe's, spaces, and extra punctuation


Tests for q3_1 (points: None)
```
>>> zafon[6]
'i'
```
failure_message: Check that you are using the correct key.


Tests for q4_1 (points: None)
```
>>> number_of_valid_values == 624
True
```
failure_message: Check that you are returning and not printing.


Tests for q4_2 (points: None)
```
>>> mult_inverse(17)
23
```
failure_message: Check that you are using the egcd function correctly in your function.


Tests for q4_3 (points: None)
```
>>> big("A",17,1)
'd'
```
failure_message: Check that you are using the egcd function correctly in your function.
```
>>> big("HXZDMZiXP",17,1)
'imgutgfmK'
```
failure_message: Check that you are subtracting the additive key first. Then check that you are moding by the length of the big alphabet.
