```python
# Initialize Otter
import otter
grader = otter.Notebook("lab03.ipynb")
```

# Lab 03: Caesar Cipher with Larger Alphabets

In today's lab, you'll learn how to implement Caesar cipher using alphabets with more than 26 characters.

This lab covers parts of [Chapter 4](https://macs4200.org/chapters/04/2/caesar-cipher.html) and [Chapter 5](https://macs4200.org/chapters/05/3/file-io.html) of the online textbook. Let's get started!

## Big Alphabets

In this section you'll complete the Caesar cipher on alphabets that contain more than 26 characters.

Suppose we extend the alphabet to include uppercase letters and lowercase letters. We'll now have 52 characters to work with. Let's agree that capital letters will come first and use the positional numbers of 0 - 25, and then lowercase letters will come second and use the position numbers 26 - 51.

```python
BIG_LETTERS = 'ABCDEFGHIJKLMNOPQRSTUVWXYZabcdefghijklmnopqrstuvwxyz'
```

Remember, to Python 'Q' and 'q' are two different characters. They'll have different numerical values in our system as well.

```python
print( BIG_LETTERS[16] )
```

```python
print( BIG_LETTERS[42] )
```

### Question 1.1

Below you'll find similar code used in an earlier problem. You'll even find the same plaintext message. However, since we now want to consider uppercase and lowercase letters as valid, you'll need to change the way you clean the text. Additionally, the alphabet is now larger. What else will you need to change to ensure that all 52 possible plaintext letters are mapped to 52 different ciphertext letters.

**Note:** Don't forget, the alphabet is now stored to `BIG_LETTERS`, not `LETTERS` so adjust accordingly!

**Note #2:** Don't forget to clean the message. It should not contain spaces, exclamation points, periods, question marks, or commas.

```python
BIG_LETTERS = 'ABCDEFGHIJKLMNOPQRSTUVWXYZabcdefghijklmnopqrstuvwxyz'
plaintext = 'What is the answer, you ask? I will tell you!'
clean_plaintext = ...
key = 37 # Don't change this key or it will break the test

# Creating an empty string that will eventually store your ciphertext
ciphertext = ''

for char in clean_plaintext:
    plaintext_numerical = ...
    ciphertext_numerical = ...
    ciphertext = ...
    
print(ciphertext)
```

```python
grader.check("q1_1")
```

It's important to realize that we could define the alphabet to contain as many characters as we want! For most messages up through the introduction of computers, English messages contained only the 26 character alphabet we've been using so far. However, there's no mathematical reason why the alphabets couldn't be larger or smaller, as long as we tend to modulus of our calculations. We'll explore other sized alphabets throughout the course.

Now that you've got a good handle on implementing the Caesar cipher for various size alphabets, let's turn our attention to formatting the ciphertext. Ciphertext is conventionally printed grouped into blocks of length 5, so let's see if we can figure out how to take an unblocked message and group it appropriately. 

### `if` statements

To help out with this task, you may find `if` statements to be helpful. These allow you to run a line of code only **`if`** a certain criteria is met. Run the example below for different values of x and take note of the structure of the statement.

```python
x = 5
if x > 3:
    print('x is larger than 3')
```

The `print` statement is only run with the numerical value stored to `x` meets the criteria of `x > 3`. If `x` does not meet that criteria, then no additional code is run.

### Question 2.2

Below you'll find some starter code that will iterate over the string `unformatted`. You should add to this code so that when completed the string `formatted` contains the same characters as `unformatted` but grouped into blocks of length 5. If the last block doesn't have 5 characters, that's okay!

For this code, try using the following strategy inside the loop:
* append `char` to the `formatted` string
* check to see if `formatted` has a multiple of 5 characters in it (not counting spaces)
  * if it does, also append a space
  * if it doesn't, do nothing!

```python
unformatted = 'HELLOTHISISFORTESTING'
formatted = ''

for char in unformatted:
    ...

print(formatted)
```

```python
grader.check("q1_2")
```

# Submitting your work
You're done with this Lab! All assignments in the course will be distributed as notebooks like this one, and you will submit your work by doing the following:
* Save your notebook
* Restart the kernel and run up to this cell.
* Run all the tests by running the cell containing `grader.check_all()`. Make sure they pass the way you expect them to.
* Run the cell below with the code `grader.export(...)`.
* Download the file named `labXX.zip`, found in the explorer pane on the left side of the screen.
* Upload `labXX-<date-time stamp>.zip` to the corresponding lab assignment on Canvas.

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit.

When done exporting, find the `.zip` file in the left side of the screen in the file browser, right-click, and select **Download**. You'll submit this `.zip` file for the assignment Gradescope for grading.

```python
grader.export(pdf=False, force_save=True)
```

Tests for q1_1 (points: 1)
```
>>> any( bad_chars in clean_plaintext for bad_chars in ['?', '!', ' ', '.', ','] )
False
```
failure_message: You still have some punctuation in your message after cleaning
```
>>> ciphertext
'HSLeTdeSPLYdhPcjZfLdVthTWWePWWjZf'
```


Tests for q1_2 (points: 1)
```
>>> formatted.count(' ')
4
```
failure_message: There should be 4 spaces (5 blocks) in this message
```
>>> len(formatted)
25
```
failure_message: There should be 25 total characters in the string (including spaces)
```
>>> formatted
'HELLO THISI SFORT ESTIN G'
```
